---
title       : The Whys and Hows of Tidy Data
description : Why we need tidy data and using `tidyr` to make messy data tidy


--- type:NormalExercise lang:r xp:0 skills:1 key:2b27344379
## What you'll learn in this Chapter:

+ How to load data
+ How to identify tidy data
+ How to make messy data tidy using `tidyr`
+ Long vs. Wide data
+ The tidy manifesto
+ Joining tables on a common identifier

After this chapter, you will be able to handle about 90% of your data visualization,
data manipulation, and data cleaning problems. Hopefully you'll have gained enough
confidence to start going beyond this Bootcamp and learn some more.

*** =instructions
Just move on to the next exercise.

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:24146e724e
## Loading data

We've been hiding some of the complexity of the analysis procedure from you until now. One of the 
most confusing parts about R to learn is how to load data. We'll take it step by step.

There are three main formats we'll have to deal with:

+ Comma Separated Values (`.csv` files) - use the `readr` package
+ Tab Delimited Files (`.tsv` or `.txt` files) - use the `readr` package
+ Excel Spreadsheets (`.xls` and `.xlsx` files) - use the `readxl` package

The functions in these packages will make a guess as to the formatting of the columns, but
sometimes you will have to fix these. We'll show you how to do that in the next section.

*** =instructions


*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```


--- type:NormalExercise lang:r xp:100 skills:1 key:efed4e694f
## Long versus Wide Data


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```


--- type:NormalExercise lang:r xp:100 skills:1 key:87aa2a11b4
## Joining two tables together


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
