---
title       : Simple Stats and Modeling with broom.
description : Now we have tidy data, let's start doing some statistics!


--- type:NormalExercise lang:r xp:100 skills:1 key:e91efd94b6
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
