---
title       : Statistics with Tidy Data
description : Statistics made easier with Tidy Data


--- type:NormalExercise lang:r xp:100 skills:1 key:e91efd94b6
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
